# CSCI 3155 Recitation 8
March 15

### Examples

https://docs.google.com/presentation/d/1OlAu7pjBfKogQg6Z5Ykkid9FfWVH-yrSbPUCHMI-3HM/edit?usp=sharing

## Exercise
Given the following functions, change them to use continuations:

```scala
def doThing1(): Int = 5 + 3
def doThing2(a: Int, b: Int): Int = a + b
def doThing3(f: Int => String): Int = {
    println(f(5))
    5
}
```

Output:
```
defined [32mfunction[39m [36mdoThing1[39m
defined [32mfunction[39m [36mdoThing2[39m
defined [32mfunction[39m [36mdoThing3[39m
```

```scala
// BEGIN SOLUTION
def doThing1_cont(k: Int => Int): Int = k(5 + 3)
def doThing2_cont(a: Int, b: Int, k: Int => Int): Int = k(a + b)
def doThing3_cont(f: Int => String, k: Int => Int): Int = {
    println(f(5))
    k(5)
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mdoThing1_cont[39m
defined [32mfunction[39m [36mdoThing2_cont[39m
defined [32mfunction[39m [36mdoThing3_cont[39m
```

## Exercise
Update the following expression to use the new continuation functions:

```scala
val x = doThing1()
val y = doThing2(x, 3)
val z = doThing3(n => "Log: " + n.toString)
```

Output:
```
Log: 5
[36mx[39m: [32mInt[39m = [32m8[39m
[36my[39m: [32mInt[39m = [32m11[39m
[36mz[39m: [32mInt[39m = [32m5[39m
```

```scala
// BEGIN SOLUTION
val z2 = 
    doThing1_cont(x =>
        doThing2_cont(x, 3, y =>
            doThing3_cont(n => "Log: " + n.toString, n =>
                n)))
// END SOLUTION

assert(z == z2)
```

Output:
```
Log: 5
[36mz2[39m: [32mInt[39m = [32m5[39m
```

### Exercise: Backtracking
As with any good programing tool, we can use continuations to solve problems without putting them everywhere in our code. In this example, you will write a search function that looks for a value in a binary tree, **not** a B**S**T. Take advantage of continuations to remember where to look if something isn't found while keeping the function tail recursive.

<!-- grade_id: cell-bdaf875cad5b0678, points: 0, grade: false, solution: true -->
```scala
sealed trait Tree
case object Empty extends Tree
case class Node(l: Tree, d: Int, r: Tree) extends Tree

def search(t: Tree, i: Int, fail_continuation: () => Boolean): Boolean =
    // BEGIN SOLUTION
    t match {
        case Empty => fail_continuation()
        case Node(l, j, r) if i == j =>
            true
        case Node(l, j, r) =>
            search(l, i, () => {
                search(r, i, () => {
                    fail_continuation()
                })
            })
    }
    //END SOLUTION

```

Output:
```
defined [32mtrait[39m [36mTree[39m
defined [32mobject[39m [36mEmpty[39m
defined [32mclass[39m [36mNode[39m
defined [32mfunction[39m [36msearch[39m
```

<!-- grade_id: cell-df2188387bae0d26, points: 1, grade: true, solution: false -->
```scala
val t = Node(Empty, 10, Node(Empty, 6, Empty))
assert(search(t, 10, () => false))
assert(!search(t, 0, () => false))
```

Output:
```
[36mt[39m: [32mNode[39m = Node(Empty,10,Node(Empty,6,Empty))
```

### Exercise: Eval (again)

This is similar to the example in class: implement eval for the small language given below, ensuring that every function call is a tail call.

<!-- grade_id: cell-83c17e7eeba4b4e4, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr
case class BoolLiteral(b: Boolean) extends Expr
case class And(left: Expr, right: Expr) extends Expr
case class If(test: Expr, then: Expr, otherwise: Expr) extends Expr

def eval(e: Expr, continuation: Boolean => Boolean): Boolean =
    // BEGIN SOLUTION
    e match {
        case BoolLiteral(b) =>
            continuation(b)
        case And(left, right) =>
            eval(left, left_val =>
                 eval(right, right_val =>
                      continuation(left_val && right_val)
                     )
                )
        case If(test, then, otherwise) =>
            eval(test, test_val =>
                 if (test_val) eval(then, b => continuation(b))
                 else eval(otherwise, continuation)
                )
    }
    // END SOLUTION
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mBoolLiteral[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mIf[39m
defined [32mfunction[39m [36meval[39m
```

<!-- grade_id: cell-c127bcc04c6197cc, points: 1, grade: true, solution: false -->
```scala
val e = If(
    And(BoolLiteral(true), BoolLiteral(true)),
    BoolLiteral(false),
    BoolLiteral(true)
)
assert(!eval(e, x => x))
```

Output:
```
[36me[39m: [32mIf[39m = If(And(BoolLiteral(true),BoolLiteral(true)),BoolLiteral(false),BoolLiteral(true))
```

## Exercise: Async Continuations (Callbacks)
Use the `http_get` function to request the given url and provide a continuation that will append `" lie."` to the result

```scala
import $ivy.`com.lihaoyi::requests:0.1.4`
import scala.concurrent.{ Future, Await }
import scala.concurrent.duration._
import scala.concurrent.ExecutionContext.Implicits.global

def http_get(url: String, k: String => String): Future[String] =
    Future(requests.get(url)).map(_.text).map(k)
```

Output:
```
[32mimport [39m[36m$ivy.$                            
[39m
[32mimport [39m[36mscala.concurrent.{ Future, Await }
[39m
[32mimport [39m[36mscala.concurrent.duration._
[39m
[32mimport [39m[36mscala.concurrent.ExecutionContext.Implicits.global

[39m
defined [32mfunction[39m [36mhttp_get[39m
```

```scala
val shakira_url = "http://www.mocky.io/v2/5c6fa3fc3400004e5f8931a6"

// BEGIN SOLUTION
val res = http_get(shakira_url, result => result + " lie.")
// END SOLUTION
```

Output:
```
[36mshakira_url[39m: [32mString[39m = [32m"http://www.mocky.io/v2/5c6fa3fc3400004e5f8931a6"[39m
[36mres[39m: [32mFuture[39m[[32mString[39m] = Future(<not completed>)
```

```scala
// Run this cell if you're using the older scala kernel (from VM)
// which doesn't update the future
Await.result(res, 5000 millis)
```

Output:
```
[36mres10[39m: [32mString[39m = [32m"Hips don't lie."[39m
```